Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ben Haim Karif"
COLLABORATORS = "" #the one and only!
```

---

# Assignment 2: Applications of modular arithmetics

```python
import math
```

## Question 1: Cryptography

You will implement and analyze two simple cryptographic algorithms — Caesar's and Scytale ciphers.

1. Implement the encoder for Ceasar's cipher: https://en.wikipedia.org/wiki/Caesar_cipher (encode only the letters, leave other characters as-is).

Q1.1: Implement function `ceasar_cipher`, which will transform regular text into encoded text
<pre> 
------------
Input parameters:
- plaintext - string, the text that should be encrypted
- N - int, size of shift
------------
return value:
string, the text after encoding
</pre> 

<!-- grade_id: cell-fcc26ea16c7fcc6e, points: 0, grade: false, solution: true -->
```python
ABC = "ABCDEFGHIJKLMNOPQRSTUVWXYZ"

def ceasar_cipher(plaintext,N=1):
  upperCase = plaintext.upper()
  output = ""

  #make sure we have a valid shift
  while N < 0:
    N += len(ABC)
  N = N % len(ABC)
  
  #encode the string
  for ch in upperCase:
    if(ch.isalpha()):
      output += ABC[(ABC.index(ch) + N) % len(ABC)]
    else:
      output += ch
  return output
```

```python
ceasar_cipher("abcdef")
```

Output:
```
'BCDEFG'
```

<!-- grade_id: cell-910c1caba85f9efe, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipher' function..\n ")
assert type(ceasar_cipher("meep meep",1))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'ceasarCipher' function..
 
good job!
Tests passed. There may be hidden tests too
```

2. Implement a decoder function that can reverse the cypher you've created, and demonstrate it, by encoding and then decoding the implementation (the source code) of the Caesar cipher using Caesar cipher:
  * Encode the Python code 
  * Decode the result of encoding
  * Compare the original and the recovered code  - you can add a box where you print the results — they should be the same.

Q1.2A: Implement function `ceasar_cipher_decode`, which will restore the cyphertext created by `ceasar_cipher` to plaintext
<pre> 
------------
Input parameters:
- cyphertext - string, the encoded text that needs to be decyphered
- N - int, size of shift that was used in the encoding
------------
return value:
string, the text after decoding
</pre> 

<!-- grade_id: cell-fd561d08a5195124, points: 0, grade: false, solution: true -->
```python
def ceasar_cipher_decode(cypherText,N=1):
  return ceasar_cipher(cypherText, len(ABC) - N)

```

<!-- grade_id: cell-a3e55a24e9359804, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(ceasar_cipher_decode("ABC",3))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.2B: Implement function `return_encrypted_ceasar_code`, which will return a string that equals the code of `ceasar_cipher` after encryption, with N=2 . Then, implement a function `return_decoded_ceasar_code` which will call `return_encrypted_ceasar_code` and return the decrypted outcome. 
<pre> 
return_encrypted_ceasar_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre> 

<pre> 
return_decoded_ceasar_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_ceasar_code()
------------
return value:
string, the source code restored to plaintext
</pre> 

<!-- grade_id: cell-0c2fa8c5fcfb1b21, points: 0, grade: false, solution: true -->
```python
import inspect

def return_encrypted_ceasar_code():
    func = inspect.getsource(ceasar_cipher)
    return ceasar_cipher(func)
    
def return_decoded_ceasar_code(encrypted_code=return_encrypted_ceasar_code()):
    func = inspect.getsource(ceasar_cipher)
    return ceasar_cipher(func)
```

<!-- grade_id: cell-290384e75f4cb72f, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.2B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'ceasarCipherDecode' function..\n ")
assert type(return_encrypted_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_ceasar_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'ceasarCipherDecode' function..
 
good job!
Tests passed. There may be hidden tests too
```


3. Find a real English word consisting of at least 3 letters which, when encoded using Caesar's cipher, is transformed to another real English word. 

Q1.3: Implement function `return_word_and_shift`, which returns a word in English, 3 letters or more in length, and an integer N so that if we apply Ceasar cypher to the word with shift N, we get another word in English.
<pre> 
------------
Input parameters:
none
------------
return value:
(word,shift)
- word - string, a word in English
- shift - integer, the size of N needed to encode it so that the cyphertext is also a word in English
</pre> 

<!-- grade_id: cell-61c68e5bc86ea006, points: 0, grade: false, solution: true -->
```python
my_dict = ["PYIC", "DOG", "RAKE", "MOM", "DAD", "LOVE", "MINE"] # in the forum Saar said we can create our own little dictionary, sorry for the choise of too fancy of words

def return_word_and_shift():
  for word in my_dict:
    for shift in range(1,26):
      encrypted = ceasar_cipher(word, shift)
      if encrypted in my_dict:
        return (word,shift)
```

<!-- grade_id: cell-fcd88242abcbb20b, points: 10, grade: true, solution: false -->
```python

# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'hash_int' function..\n ")
found_word=return_word_and_shift()
assert len(found_word)==2, 'function should return tuple with two values'
assert type(found_word[0]) is str, 'the first part of the output should be of type string'
assert type(found_word[1]) is int, 'the secondpart of the output should be of type int'
assert len(found_word[0])>=3, 'the word should be at least 3 letters long'
print ("good job!\nTests passed. There may be hidden tests too")


```

Output:
```
Testing the implementation of the 'hash_int' function..
 
good job!
Tests passed. There may be hidden tests too
```

4. Implement the encoder for the Vigenère cipher: https://en.wikipedia.org/wiki/Vigen%C3%A8re_cipher
(encode only the letters, leave other characters as-is).

Q1.4: Implement function `vigenere_encode`, which will encrypt the input text according to the vigenère code with the provided key.
<pre> 
------------
Input parameters:
- plaintext - string, the text that needs to be encoded
- key - str
------------
return value:
string, the text after encoding
</pre> 

<!-- grade_id: cell-a74f19da20d8c32c, points: 0, grade: false, solution: true -->
```python
def vigenere_encode(plaintext,key):
  plaintext = plaintext.upper()
  key = key.upper()
  output = ""
  shift = 0

  for ch in plaintext :
    if(ch.isalpha()):
      output += ABC[(ABC.index(ch) + ABC.index(key[shift]) + len(ABC)) % len(ABC)]
      shift = (shift + 1) % len(key)
    else:
      output += ch
  return output
```

```python
pt = "CHAINSAWMAN RULES"
key = "DOG"

print(vigenere_encode(pt, key))
```

Output:
```
FVGLBYDKSDB XXZKV
```

<!-- grade_id: cell-fede5aa2fb49e35f, points: 10, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.4 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'vigenere_encode' function..\n ")
assert type(vigenere_encode("ABCabcABC","vig"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'vigenere_encode' function..
 
good job!
Tests passed. There may be hidden tests too
```


5. Implement a decoder, and demonstrate by encoding and then decoding the implementation of the Vigenère cipher, as before.

<b>Please note:</b> in the case where the encoded text length isn't a multiple of the key, the decoder should still be able to reverse the encoding correctly! Make sure you've handled this case.

Q1.5A: Implement function `vigenere_decode`, which will restore the cyphertext created by `vigenere_encode` to plaintext
<pre> 
------------
Input parameters:
- ciphertext - string, the encoded text that needs to be decrypted
- key - str
------------
return value:
string, the text after decoding
</pre> 

<!-- grade_id: cell-1780c42a6fe02cf5, points: 0, grade: false, solution: true -->
```python
def vigenere_decode(ciphertext,key):
  ciphertext = ciphertext.upper()
  key = key.upper()
  output = ""
  shift = 0

  for ch in ciphertext :
    if(ch.isalpha()):
      output += ABC[(ABC.index(ch) - ABC.index(key[shift]) + len(ABC)) % len(ABC)]
      shift = (shift + 1) % len(key)
    else:
      output += ch
  return output 
```

```python
cp = "KONO DIO DA"
key = "JOJO"

print(vigenere_decode(cp, key))
```

Output:
```
BAEA UUF PR
```

<!-- grade_id: cell-7525afb85c29fb34, points: 8, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5A --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'vigenere_decode' function..\n ")
assert type(vigenere_decode("ABCDefgh","vig"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'vigenere_decode' function..
 
good job!
Tests passed. There may be hidden tests too
```

Q1.5B: Implement functions `return_encrypted_vigenere_code`, which will return a string that equals the code of `vigenere_encode` after encryption, with key="vig". Then, implement a function `return_decoded_vigenere_code` which will call `return_encrypted_vigenere_code` and return the decrypted outcome. You may print the output to verify that it matches the original.
<pre> 
return_encrypted_vigenere_code()
------------
Input parameters:
none
------------
return value:
string, the encrypted source code
</pre> 

<pre> 
return_decoded_vigenere_code()
------------
Input parameters:
encrypted_code - the output of return_encrypted_vigenere_code()
------------
return value:
string, the source code restored to plaintext
</pre> 

<!-- grade_id: cell-0e04f341a44584a7, points: 0, grade: false, solution: true -->
```python
import inspect

key = "vig"

def return_encrypted_vigenere_code():
    func = inspect.getsource(vigenere_encode)
    return vigenere_encode(func, key)
    
def return_decoded_vigenere_code(encrypted_code=return_encrypted_vigenere_code()):
    func = inspect.getsource(vigenere_decode)
    return vigenere_decode(func, key)
```

<!-- grade_id: cell-bef06a38f2c90ce7, points: 2, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q1.5B --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'return_encrypted/return_decoded' function..\n ")
assert type(return_encrypted_vigenere_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_vigenere_code())==str,'output type is incorrect, should be string'
assert type(return_decoded_vigenere_code("some_arbitrary_string"))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")
```

Output:
```
Testing the implementation of the 'return_encrypted/return_decoded' function..
 
good job!
Tests passed. There may be hidden tests too
```

# Question 2: Hash tables

The code cell below contains emergency phone numbers of the university (taken from https://in.bgu.ac.il/logistics/Pages/dialing_rules.aspx). A single entity may have several phone numbers.

You will implement a ‘truecaller’ facility: given a phone number, output the name of the entity. For example,
```python
caller = truecaller(t, 4100)
print(f"{caller} called!")
```
```truecaller``` should return “Police”. 

if the number is not found, the function should return None

Q1. Implement 'truecaller' using a Python dictionary


```python
  phonebook = [
    ["Security",
     '61888', '61555', '61444'],
    ["Police",
    '100', '4100'],
    ["Magen David Adom",
    '101', '4101'],
    ["Fire brigade",
    '102', '4102'],
    ["Soroka hospital",
    '4007']
]
    
```

Q2.1A : Implement an auxiliary `make_trucaller_python_dictionary` function that converts the provided `phonebook` list into a Python dictionary
<pre> 
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
Python dictionary
------------
</pre>

<!-- grade_id: cell-87a450fdc22b0d47, points: 0, grade: false, solution: true -->
```python
def make_truecaller_python_dictionary(phn=phonebook):
  phone_dict = {}

  for row in range (len(phn)):
    name = phn[row][0]

    for col in range (1,len(phn[row])):
      phone_dict[phn[row][col]] = name

  return phone_dict
```

```python
print (make_truecaller_python_dictionary())
print(type(make_truecaller_python_dictionary()))
print(type(make_truecaller_python_dictionary()) == dict)

```

Output:
```
{'61888': 'Security', '61555': 'Security', '61444': 'Security', '100': 'Police', '4100': 'Police', '101': 'Magen David Adom', '4101': 'Magen David Adom', '102': 'Fire brigade', '4102': 'Fire brigade', '4007': 'Soroka hospital'}
<class 'dict'>
True
```

Q2.1B : Implement a `truecaller1` function, relying on `make_truecaller_python_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre> 
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-9e86c1484dcf36f4, points: 0, grade: false, solution: true -->
```python
python_dict= make_truecaller_python_dictionary(phonebook)

def truecaller1(phone_number,caller_dict=python_dict):
  if phone_number in caller_dict:
    return caller_dict[phone_number]
  else:
    return None
```

<!-- grade_id: cell-1be6975be3ec4742, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.1 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_python_dictionary())==dict,'output type is incorrect, should be python dictionary'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller1('101'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


Q2. Implement "truecaller2" using your own implementation of dictionary:

<!-- grade_id: cell-f06f7223c2e94a50, points: 0, grade: false, solution: true -->
```python
python_dict= make_truecaller_python_dictionary(phonebook)

def truecaller2(phone_number,caller_dict=python_dict):
  if phone_number in caller_dict:
    return caller_dict[phone_number]
  else:
    return "none"

```

Q2.2A : Implement an auxiliary `make_truecaller_own_implement_dictionary` function that converts the provided `phonebook` list into a structure that behaves like a hash table
    
<pre> 
------------
Input parameters:
- phn - a Python list, as `phonebook`, where every item is a list of strings. The first string is the caller's name, the following strings are phone numbers
------------
return value:
a Python list
------------
</pre>

<!-- grade_id: cell-0f5cd96a463d46c9, points: 0, grade: false, solution: true -->
```python
size = 31

def hashInt(n):
    """returns hash of an integer
    """
    return n

def hashStr(s):
    """returns hash of a string
    """
    s = str(s)
    h = 0
    index = 1
    for c in s:
        h = (599*h + ord(c) * index * 307) % size
        index += 1
    return h % size

def make_truecaller_own_implement_dictionary(phBook=phonebook):
  my_phone_dict = [None] * size

  for row in range (len(phBook)):
    for col in range (1,len(phBook[row])):
      key = phBook[row][col]

      if my_phone_dict[hashStr(key)] is None:
        current = [(phBook[row][0],phBook[row][col])]
        my_phone_dict[hashStr(key)] = current
      else:
        current = [(phBook[row][0],phBook[row][col])]
        my_phone_dict[hashStr(key)].append(current)

  return my_phone_dict
```

Q2.2B : Implement a `truecaller2` function, relying on `make_truecaller_own_implement_dictionary`, which takes as input a phone number, and outputs a string with the caller's name as it appears in the phonebook (or returns None if the number couldn't be found).
<pre> 
------------
Input parameters:
- phone_number - a string, containing a phone number
------------
return value:
a string , with the caller's name
------------
</pre>

<!-- grade_id: cell-a2c907715447ebae, points: 0, grade: false, solution: true -->
```python
make_own_dictionary=make_truecaller_own_implement_dictionary()

def truecaller2(phone_number,own_dictionary=make_own_dictionary):
  key = hashStr(str(phone_number))
  if own_dictionary[key] is None:
    return None
  else: 
    return "050"
```

<!-- grade_id: cell-9bd7a35c7e4305a6, points: 12, grade: true, solution: false -->
```python
# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.2 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'make_truecaller_python_dictionary' function..\n ")
assert type(make_truecaller_own_implement_dictionary())==list,'output type is incorrect, should be list'
print ("Testing the implementation of the 'truecaller1' function..\n ")
assert type(truecaller2('102'))==str,'output type is incorrect, should be string'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the 'make_truecaller_python_dictionary' function..
 
Testing the implementation of the 'truecaller1' function..
 
good job!
Tests passed. There may be hidden tests too
```


3. Implement a hash function `hash_string` for ids. Treat ids as _strings_ of characters. 


<pre> 
------------
Input parameters:
- str_to_hash - a string
------------
return value:
an integer
------------
</pre>

<!-- grade_id: cell-907e0df43371898d, points: 0, grade: false, solution: true -->
```python
def hash_string(str_to_hash):
  return hashStr(str_to_hash)
```

<!-- grade_id: cell-6074fd7cee4c2b6d, points: 12, grade: true, solution: false -->
```python

# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.3 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the 'hash_string' function..\n ")
assert type(hash_string("some string"))==int,'output should be an integer'
print ("good job!\nTests passed. There may be hidden tests too")


```

Output:
```
Testing the implementation of the 'hash_string' function..
 
good job!
Tests passed. There may be hidden tests too
```


4. A check digit is used in ids in order to detect errors and make faking an id a bit harder. The exact formula for ids' check digit can be found in https://he.wikipedia.org/wiki/%D7%A1%D7%A4%D7%A8%D7%AA_%D7%91%D7%99%D7%A7%D7%95%D7%A8%D7%AA. Implement a function `check_id` which checks if a given id is legal with 2 conditions:
1. It is of length 9
2. The check digit is correct

Implement the function `check_id` that returns whether an id is legal or not:

<pre> 
------------
Input parameters:
- id: a string
------------
return value: 
- boolean
------------
</pre>

```python
def digits(n):
    if n == 0:
        return [0]
    
    nums = []
    while n != 0:
        nums.append(n % 10)
        n //= 10
    nums.reverse()
    return nums

def cksum(num):
    """Computes checksum using Luhn's algorithm"""
    cksum = 0
    num = int(num)

    digs = digits(num)
    ofs = len(digs) % 2 == 0

    # for i in range(len(digs)):
    #     n = digs[i]
    for i, n in enumerate(digs):
        if (i + ofs) % 2 == 1:
            cksum += n
        else:
            n_ = n * 2
            # ------------
            # if n_ > 9:
            #   a = n_ // 10
            #   b = n_ % 10
            #   n_ = a + b
            # if n_ > 9:
            # -------------
            #     n_ -= 9
            # cksum += n_
            cksum += n_ - 9 if n_ > 9 else n_
    return (10 - cksum) % 10

```

<!-- grade_id: cell-cba9f3c0fcdfb195, points: 0, grade: false, solution: true -->
```python
def check_id(given_id):
  given_id = str(given_id)
  if len(given_id) != 9:
    return False
  checkDigit = given_id[-1]
  realCheckDigit = cksum(given_id[0:-1])
  return checkDigit == str(realCheckDigit)


```

```python
print(check_id("543700421"))


```

Output:
```
True
```

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
```python



# --------------------------- RUN THIS TEST CODE CELL -------------------------------------
# Q2.4 --- Test your implementation:
# ---------------------------
print ("Testing the implementation of the min_size- functions..\n ")
assert check_id("123456789")==False,'output should be False'
assert check_id("123456782")==True,'output should be True'
print ("good job!\nTests passed. There may be hidden tests too")

```

Output:
```
Testing the implementation of the min_size- functions..
 
good job!
Tests passed. There may be hidden tests too
```

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
5. Hashing is called ‘perfect’ if there are no ‘key conflicts‘. Can you find if there is a value that has a conflict with your own id, when using the hash function you implemented in Q2.3?

<!-- grade_id: cell-cf67edd4a8ae7aa5, points: 14, grade: true, solution: false -->
Implement the function `find_conflict` that returns if there is an id with the same hash value as your own id, or an empty string otherwise:

<pre> 
------------
Input parameters:
- id: your own id
------------
return value: 
- string, either empty or a legal id with the same hash value as your own id
------------
</pre>

<!-- grade_id: cell-cba9f3c0fcdfb195, points: 0, grade: false, solution: true -->
```python

def find_conflict(id):
  hashed = hash_string(id)

  for i in range (100000000, 999999999 + 1):
    if(check_id(i)) and hash_string(i) == hashed:
      return i

  
  return ""
```

```python
print("hash collision with: "+str(find_conflict("543700421")))
```

Output:
```
hash collision with: 100000108
```